# Assignment 2
For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```


<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np


def proportion_of_education():
    df = pd.read_csv("assets/NISPUF17.csv", index_col = 0)
    print(df)

    LOE = df['EDUC1']  #find the column relating to the maternity age
    n = len(LOE)    #Total number of entries in the data

    '''
    or
    
    test = LOE.value_counts(normalize=True)    #.value_counts gives the proprtion of items present in the result
    print(test)
    print(test.iloc[3])

    '''
    A = np.sum(LOE ==1)/n #Total number of entries with 1 as the value divided by the total number of entries
    B = np.sum(LOE ==2)/n
    C = np.sum(LOE ==3)/n
    D = np.sum(LOE ==4)/n
    
    loe = {"less than high school": A,
    "high school": B,
    "more than high school but not college": C,
    "college": D
    }

    print(loe)
    
    return loe

   

proportion_of_education()
```

Output:
```
SEQNUMC  SEQNUMHH  PDAT    PROVWT_D     RDDWT_D  STRATUM  YEAR  \
1       128521     12852     2         NaN  235.916956     1031  2017   
2        10741      1074     2         NaN  957.353840     1068  2017   
3       220011     22001     2         NaN  189.611299     1050  2017   
4        86131      8613     1  675.430817  333.447418     1040  2017   
5       227141     22714     1  482.617748  278.768063     1008  2017   
...        ...       ...   ...         ...         ...      ...   ...   
28461    19141      1914     2         NaN   21.188088     2062  2017   
28462   121401     12140     2         NaN   26.503010     2004  2017   
28463   128831     12883     2         NaN   27.520652     2004  2017   
28464    82921      8292     2         NaN   22.390587     2062  2017   
28465   244621     24462     2         NaN   27.520652     2004  2017   

       AGECPOXR  HAD_CPOX  AGEGRP  ...  XVRCTY2  XVRCTY3  XVRCTY4  XVRCTY5  \
1           NaN         2       1  ...              
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
def average_influenza_doses():
    df = pd.read_csv("assets/NISPUF17.csv", index_col = 0)
    
    BF = df[df['CBF_01'] == 1] #filters data based on breastfed children
    #df.dropna
    NF = df[df['CBF_01'] == 2] #filters data based on non-breastfed children

    print(BF["P_NUMFLU"].mean())

    return (BF["P_NUMFLU"].mean(), NF["P_NUMFLU"].mean())

    '''
    nB = len(BF) #computes the number of breastfed children
    nF = len(NF) #computes the number of non-breastfed children

    A = BF['P_NUMFLU']
    B = NF['P_NUMFLU']
    
    print(A.sum(axis=0))
    print(B.sum(axis=0))

    RATIO = (NUM_BF/nB , NUM_NF/nF) 

    print(len(df))

    #NUM_BF = np.sum(df['P_NUMFLU']) #sums the number of doses available to breastfed children
    #NUM_NF = np.sum(df['P_NUMFLU']) #sums the number of doses available to non-breastfed children

    #RATIO = (nB/NUM_BF , nF/NUM_NF) #returns a tuple of the average number of vaccines available per child


    #return RATIO  
    '''
    
  
    
    
average_influenza_doses()
```

Output:
```
1.8799187420058687
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    # YOUR CODE HERE
    df = pd.read_csv("assets/NISPUF17.csv", index_col = 0)

    HAD_VRC = df[df['P_NUMVRC'] > 0]  #creates a dataframe containing vaccinated children
  
    HAD_CPOX = HAD_VRC[df['HAD_CPOX'] == 1] #creates a dataframe from HAD-VRC containing vaccinated children that had chickenpox
   
    HAD_NPOX = HAD_VRC[df['HAD_CPOX'] == 2 ] #creates a dataframe from HAD_VRC containing vaccinated children that did not have chickenpox
   


    #FOR MALE
    MALE_HAD_CPOX = len(HAD_CPOX[df['SEX']== 1]) 
    MALE_HAD_NPOX = len(HAD_NPOX[df['SEX']== 1])

    A = MALE_HAD_CPOX/MALE_HAD_NPOX

    #FOR FEMALE
    FEMALE_HAD_CPOX = len(HAD_CPOX[df['SEX']== 2])
    FEMALE_HAD_NPOX = len(HAD_NPOX[df['SEX']== 2])

    B = FEMALE_HAD_CPOX/FEMALE_HAD_NPOX

    dict = { "male" : A, "female" : B}

    print(dict)

    

chickenpox_by_sex()
```

Output:
```
13995
121
13830
{'male': 0.009675583380762664, 'female': 0.0077918259335489565}
C:\Users\ayool\AppData\Local\Temp\ipykernel_23616\1302905416.py:7: UserWarning: Boolean Series key will be reindexed to match DataFrame index.
  HAD_CPOX = HAD_VRC[df['HAD_CPOX'] == 1]
C:\Users\ayool\AppData\Local\Temp\ipykernel_23616\1302905416.py:9: UserWarning: Boolean Series key will be reindexed to match DataFrame index.
  HAD_NPOX = HAD_VRC[df['HAD_CPOX'] == 2 ]
C:\Users\ayool\AppData\Local\Temp\ipykernel_23616\1302905416.py:14: UserWarning: Boolean Series key will be reindexed to match DataFrame index.
  MALE_HAD_CPOX = len(HAD_CPOX[df['SEX']== 1])
C:\Users\ayool\AppData\Local\Temp\ipykernel_23616\1302905416.py:15: UserWarning: Boolean Series key will be reindexed to match DataFrame index.
  MALE_HAD_NPOX = len(HAD_NPOX[df['SEX']== 1])
C:\Users\ayool\AppData\Local\Temp\ipykernel_23616\1302905416.py:20: UserWarning: Boolean Series key will be reindexed to match DataFrame index.
  FEMALE_HAD_CPOX = len
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no’s) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn’t really the full picture, since we are not looking at when the dose was given. It’s possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
   
    from scipy import stats
    import numpy as np
    import pandas as pd
    
    # this is just an example dataframe
    #df=pd.DataFrame({"had_chickenpox_column":np.random.randint(1,3,size=(100)),
    #               "num_chickenpox_vaccine_column":np.random.randint(0,6,size=(100))})

    # here is some stub code to actually run the correlation
    #corr, pval=stats.pearsonr(df["had_chickenpox_column"],df["num_chickenpox_vaccine_column"])
    df = pd.read_csv("assets/NISPUF17.csv", index_col = 0)
    df = df[df["HAD_CPOX"] <= 2]
  
    df = df[~df["P_NUMVRC"].isna() & ~df["HAD_CPOX"].isna()]
    corr, pval=stats.pearsonr(df["HAD_CPOX"], df["P_NUMVRC"])
    # just return the correlation
    #return corr

    # YOUR CODE HERE
    
corr_chickenpox()
```

Output:
```
SEQNUMC  SEQNUMHH  PDAT    PROVWT_D     RDDWT_D  STRATUM  YEAR  \
1       128521     12852     2         NaN  235.916956     1031  2017   
2        10741      1074     2         NaN  957.353840     1068  2017   
3       220011     22001     2         NaN  189.611299     1050  2017   
4        86131      8613     1  675.430817  333.447418     1040  2017   
5       227141     22714     1  482.617748  278.768063     1008  2017   
...        ...       ...   ...         ...         ...      ...   ...   
28461    19141      1914     2         NaN   21.188088     2062  2017   
28462   121401     12140     2         NaN   26.503010     2004  2017   
28463   128831     12883     2         NaN   27.520652     2004  2017   
28464    82921      8292     2         NaN   22.390587     2062  2017   
28465   244621     24462     2         NaN   27.520652     2004  2017   

       AGECPOXR  HAD_CPOX  AGEGRP  ...  XVRCTY2  XVRCTY3  XVRCTY4  XVRCTY5  \
1           NaN         2       1  ...              
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```

Output:
```
SEQNUMC  SEQNUMHH  PDAT    PROVWT_D     RDDWT_D  STRATUM  YEAR  \
1       128521     12852     2         NaN  235.916956     1031  2017   
2        10741      1074     2         NaN  957.353840     1068  2017   
3       220011     22001     2         NaN  189.611299     1050  2017   
4        86131      8613     1  675.430817  333.447418     1040  2017   
5       227141     22714     1  482.617748  278.768063     1008  2017   
...        ...       ...   ...         ...         ...      ...   ...   
28461    19141      1914     2         NaN   21.188088     2062  2017   
28462   121401     12140     2         NaN   26.503010     2004  2017   
28463   128831     12883     2         NaN   27.520652     2004  2017   
28464    82921      8292     2         NaN   22.390587     2062  2017   
28465   244621     24462     2         NaN   27.520652     2004  2017   

       AGECPOXR  HAD_CPOX  AGEGRP  ...  XVRCTY2  XVRCTY3  XVRCTY4  XVRCTY5  \
1           NaN         2       1  ...              
```