# Assignment 1
For this assignment you are welcomed to use other regex resources such a regex "cheat sheets" you find on the web.



Before start working on the problems, here is a small example to help you understand how to write your own answers. In short, the solution should be written within the function body given, and the final result should be returned. Then the autograder will try to call the function and validate your returned result accordingly. 

```python
def example_word_count():
    # This example question requires counting words in the example_string below.
    example_string = "Amy is 5 years old"
    
    # YOUR CODE HERE.
    # You should write your solution here, and return your result, you can comment out or delete the
    # NotImplementedError below.
    result = example_string.split(" ")
    return len(result)

    #raise NotImplementedError()
example_word_count()
```

Output:
```
5
```

## Part A

Find a list of all of the names in the following string using regex.

<!-- grade_id: names, points: 0, grade: false, solution: true -->
```python
import re
def names():
    simple_string = """Amy is 5 years old, and her sister Mary is 2 years old. 
    Ruth and Peter, their parents, have 3 kids."""

    # YOUR CODE HERE
    pattern ='[A-Z][a-z]*'
    names = re.findall(pattern,simple_string)
    return names
    #raise NotImplementedError()
```

<!-- grade_id: correct_names, points: 1, grade: true, solution: false -->
```python
assert len(names()) == 4, "There are four names in the simple_string"

```

## Part B

The dataset file in [assets/grades.txt](assets/grades.txt) contains a line separated list of people with their grade in 
a class. Create a regex to generate a list of just those students who received a B in the course.

<!-- grade_id: grades, points: 0, grade: false, solution: true -->
```python
import re
def grades():
    with open ("assets/grades.txt", "r") as file:
        grades = file.read()

    # YOUR CODE HERE

    grades = re.findall ('\n(.*(?=: B))', grades)
    

    print(grades)
    #raise NotImplementedError()
grades()
```

Output:
```
['Bell Kassulke', 'Simon Loidl', 'Elias Jovanovic', 'Hakim Botros', 'Emilie Lorentsen', 'Jake Wood', 'Fatemeh Akhtar', 'Kim Weston', 'Yasmin Dar', 'Viswamitra Upandhye', 'Killian Kaufman', 'Elwood Page', 'Elodie Booker', 'Adnan Chen', 'Hank Spinka', 'Hannah Bayer']
```

<!-- grade_id: correct_grades, points: 1, grade: true, solution: false -->
```python

assert len(grades()) == 16

```

Output:
```
c:\Users\ayool\Desktop\COURSERA\Python for Data Science\Files_2\Files\home\jovyan\work\assignments\assignment1\assignment1.ipynb Cell 9' in <cell line: 1>()
----> <a href='vscode-notebook-cell:/c%3A/Users/ayool/Desktop/COURSERA/Python%20for%20Data%20Science/Files_2/Files/home/jovyan/work/assignments/assignment1/assignment1.ipynb#ch0000008?line=0'>1</a> len(grades())
      <a href='vscode-notebook-cell:/c%3A/Users/ayool/Desktop/COURSERA/Python%20for%20Data%20Science/Files_2/Files/home/jovyan/work/assignments/assignment1/assignment1.ipynb#ch0000008?line=1'>2</a> assert len(grades()) == 16

c:\Users\ayool\Desktop\COURSERA\Python for Data Science\Files_2\Files\home\jovyan\work\assignments\assignment1\assignment1.ipynb Cell 8' in grades()
      <a href='vscode-notebook-cell:/c%3A/Users/ayool/Desktop/COURSERA/Python%20for%20Data%20Science/Files_2/Files/home/jovyan/work/assignments/assignment1/assignment1.ipynb#ch0000007?line=5'>6</a> # YOUR CODE HERE
      <a href='vscode-notebook-cell:/c%3A
```

## Part C

Consider the standard web log file in [assets/logdata.txt](assets/logdata.txt). This file records the access a user makes when visiting a web page (like this one!). Each line of the log has the following items:
* a host (e.g., '146.204.224.152') 
* a user_name (e.g., 'feest6811' **note: sometimes the user name is missing! In this case, use '-' as the value for the username.**)
* the time a request was made (e.g., '21/Jun/2019:15:45:24 -0700')
* the post request type (e.g., 'POST /incentivize HTTP/1.1' **note: not everything is a POST!**)

Your task is to convert this into a list of dictionaries, where each dictionary looks like the following:
```
example_dict = {"host":"146.204.224.152", 
                "user_name":"feest6811", 
                "time":"21/Jun/2019:15:45:24 -0700",
                "request":"POST /incentivize HTTP/1.1"}
```

<!-- grade_id: logs, points: 0, grade: false, solution: true -->
```python
import re
def logs():
    with open("assets/logdata.txt", "r") as file:
        logdata = file.read()
    
    # YOUR CODE HERE
    pattern = """
    (?P<host>.*)  #Host name
    (\ -\ ) #Separator
    (?P<user_name>.*)  #Username
    (\ \[) #Separator
    (?P<time>.*) #time
    (\]\ ") #Separator
    (?P<request>.*) #request 
    (" ) #Separator
    """
    
    log_list = []
    for item in re.finditer(pattern,logdata,re.VERBOSE):
        log_list.append(item.groupdict())
        
    return log_list
    
    #: re.findall('\d+\.\d+\.\d+\.\d+\.', logdata),
     #   'user_name': re.findall('\w[0-9]', logdata),
      #  'time': re.findall('\d+\/\[A-Z][a-z]+\/\d+\:\d+\:\d+\:\d+\s-\d+', logdata),
       # 'request': re.findall('[:upper:]\s.', logdata)  
        
        
    #raise NotImplementedError()
```

<!-- grade_id: cell-correct_logs, points: 1, grade: true, solution: false -->
```python
#print(logs())
assert len(logs()) == 979

one_item={'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1'}
assert one_item in logs(), "Sorry, this item should be in the log results, check your formating"

```

```python
len(logs())
```

Output:
```
979
```

```python
logs()[0]
```

Output:
```
{'host': '146.204.224.152',
 'user_name': 'feest6811',
 'time': '21/Jun/2019:15:45:24 -0700',
 'request': 'POST /incentivize HTTP/1.1'}
```